---
id: 60f0286404aefb0562a4fdf9
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Add a `title` element to the `head`, and give your project a title of `Registration Form`. Also, nest a self-closing `link` element in the `head` element. Give it a `rel` attribute value of `stylesheet`, a `type` attribute value of `text/css`, and an `href` attribute value of `styles.css`.

# --hints--

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

The `title` element should be within the `head` element.

```js
assert.exists(document.querySelector('head > title'));
```

Your project should have a title of `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Remember, the casing and spelling matters for the title.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Your code should have a `link` element.

```js
const link_selector = document.querySelectorAll('link');
console.log(link_selector);
assert.exists(link_selector);
```

You should not change your existing `head` tags. Make sure you did not delete the closing tag.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Your `link` element should be a self-closing element.

```js
const link_selector = document.querySelectorAll('link');
console.log(link_selector);
assert.equal(link_selector.length, 1);
//assert.match(code, /<link\s*\/?\s*?>/gmi)
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
//assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const link_element = document.querySelector('link');
console.log(link_element);
const rel_attribute = link_element.getAttribute('rel');
console.log(rel_attribute);
assert.equal(rel_attribute, "stylesheet");
```

Your `link` element should have a `type` attribute with the value `text/css`.

```js
const link_element = document.querySelector('link');
const type_attribute = link_element.getAttribute('type');
console.log(type_attribute);
assert.equal(type_attribute, "text/css");

//assert.match(code, /<link[\s\S]*?type=('|"|`)text\/css\1/)
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const link_element = document.querySelector('link');
const href_attribute = link_element.getAttribute('href');
assert.equal(href_attribute, "styles.css");
//assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  --fcc-editable-region--
  <head>
    
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
